// intro1.rs
//
// 關於這個 `I AM NOT DONE` 標註:
// 我們有時會鼓勵您在已經解決一個練習之後繼續嘗試，即使您已經找到了答案。
// 如果您已經完成了所有工作並且準備好進行下一個練習，請刪除下面的 `I AM NOT DONE` 註釋。
//
// 如果您正在使用 `rustlings watch` 運行這個文件: 每當您更改下面的一行代碼時，練習文件將會重新加載！
// 嘗試添加一行 `println!`，或者嘗試改變它在終端輸出的內容。
// 嘗試刪除一個分號，看看會發生什麼！
//
// 執行 `rustlings hint intro1` 或使用 `hint` watch 子命令來獲取提示。

// I AM NOT DONE

fn main() {
    println!("您好，");
    println!(r#"       歡迎來到...                      "#);
    println!(r#"                 _   _ _                  "#);
    println!(r#"  _ __ _   _ ___| |_| (_)_ __   __ _ ___  "#);
    println!(r#" | '__| | | / __| __| | | '_ \ / _` / __| "#);
    println!(r#" | |  | |_| \__ \ |_| | | | | | (_| \__ \ "#);
    println!(r#" |_|   \__,_|___/\__|_|_|_| |_|\__, |___/ "#);
    println!(r#"                               |___/      "#);
    println!();
    println!("此練習已成功編譯。其餘練習包含編譯或邏輯錯誤。Rustlings 的核心概念是修復這些錯誤並解決練習。祝您好運！");
    println!();
    println!("此練習的源碼在 `exercises/00_intro/intro1.rs`。請查看！");
    println!("之後，每個練習的源碼路徑將會顯示在成功或失敗的輸出結果中。");
    println!();
    println!(
        "如果您想使用 Rust 的 LSP 實現 rust-analyzer，請確保您的編輯器已設置，"
    );
    println!("然後在繼續之前運行 `rustlings lsp`。")
}

