// intro1.rs
//
// 關於這個 `I AM NOT DONE` 標註:
// 我們有時會鼓勵您在已經解決一個練習之後繼續嘗試，即使您已經找到了答案。
// 如果您已經完成了所有工作並且準備好進行下一個練習，請刪除下面的 `I AM NOT DONE` 註釋。
//
// 如果您正在使用 `rustlings watch` 運行這個文件: 每當您更改下面的一行代碼時，練習文件將會重新加載！
// 嘗試添加一行 `println!`，或者嘗試改變它在終端輸出的內容。
// 嘗試刪除一個分號，看看會發生什麼！
//
// 執行 `rustlings hint intro1` 或使用 `hint` watch 子命令來獲取提示。

// I AM NOT DONE

fn main() {
    println!("Hello and");
    println!(r#"       welcome to...                      "#);
    println!(r#"                 _   _ _                  "#);
    println!(r#"  _ __ _   _ ___| |_| (_)_ __   __ _ ___  "#);
    println!(r#" | '__| | | / __| __| | | '_ \ / _` / __| "#);
    println!(r#" | |  | |_| \__ \ |_| | | | | | (_| \__ \ "#);
    println!(r#" |_|   \__,_|___/\__|_|_|_| |_|\__, |___/ "#);
    println!(r#"                               |___/      "#);
    println!();
    println!("This exercise compiles successfully. The remaining exercises contain a compiler");
    println!("or logic error. The central concept behind Rustlings is to fix these errors and");
    println!("solve the exercises. Good luck!");
    println!();
    println!("The source for this exercise is in `exercises/00_intro/intro1.rs`. Have a look!");
    println!(
        "Going forward, the source of the exercises will always be in the success/failure output."
    );
    println!();
    println!(
        "If you want to use rust-analyzer, Rust's LSP implementation, make sure your editor is set"
    );
    println!("up, and then run `rustlings lsp` before continuing.")
}
