// variables5.rs
//
// Make me compile!
//
// If you need help, open the corresponding README.md or run: rustlings hint variables5

// I AM NOT DONE

fn main() {
    let number = "T-H-R-E-E";
    println!("Spell a Number : {}", number);
    number = 3;
    println!("Number plus two is : {}", number + 2);
}
