// variables1.rs
// Make me compile! Scroll down for hints :)

// About this `I AM NOT DONE` thing:
// We sometimes encourage you to keep trying things on a given exercise,
// even after you already figured it out. If you got everything working and
// feel ready for the next exercise, you the `I AM NOT DONE` comment below.

// I AM NOT DONE

fn main() {
    x = 5;
    println!("x has the value {}", x);
}

































// Hint: The declaration on line 5 is missing a keyword that is needed in Rust
// to create a new variable binding.
