// primitive_types3.rs
//
// Create an array with at least 100 elements in it where the ??? is.
//
// If you need help, open the corresponding README.md or run: rustlings hint primitive_types3

// I AM NOT DONE

fn main() {
    let a = ???

    if a.len() >= 100 {
        println!("Wow, that's a big array!");
    } else {
        println!("Meh, I eat arrays like that for breakfast.");
    }
}
