// primitive_types4.rs
//
// Get a slice out of Array a where the ??? is so that the test passes.
//
// If you need help, open the corresponding README.md or run: rustlings hint primitive_types4

// I AM NOT DONE

#[test]
fn slice_out_of_array() {
    let a = [1, 2, 3, 4, 5];

    let nice_slice = ???

    assert_eq!([2, 3, 4], nice_slice)
}
