// primitive_types4.rs
// Get a slice out of Array a where the ??? is so that the `if` statement
// returns true. Execute `rustlings hint primitive_types4` for hints!!

// I AM NOT DONE

#[test]
fn main() {
    let a = [1, 2, 3, 4, 5];

    let nice_slice = ???

    assert_eq!([2, 3, 4], nice_slice)
}
