// primitive_types5.rs
//
// Destructure the `cat` tuple so that the println will work.
//
// If you need help, open the corresponding README.md or run: rustlings hint primitive_types5

// I AM NOT DONE

fn main() {
    let cat = ("Furry McFurson", 3.5);
    let /* your pattern here */ = cat;

    println!("{} is {} years old.", name, age);
}
