// tests2.rs
//
// This test has a problem with it -- make the test compile! Make the test
// pass! Make the test fail!
//
// If you need help, open the corresponding README.md or run: rustlings hint tests2

// I AM NOT DONE

#[cfg(test)]
mod tests {
    #[test]
    fn you_can_assert_eq() {
        assert_eq!();
    }
}
