// variables5.rs
//
// 執行 `rustlings hint variables5` 或使用 `hint` watch 子命令來獲取提示。

// I AM NOT DONE

fn main() {
    let number = "T-H-R-E-E"; // don't change this line
    println!("Spell a Number : {}", number);
    number = 3; // don't rename this variable
    println!("Number plus two is : {}", number + 2);
}
