// strings1.rs
//
// Make me compile without changing the function signature!
//
// If you need help, open the corresponding README.md or run: rustlings hint strings1

// I AM NOT DONE

fn main() {
    let answer = current_favorite_color();
    println!("My current favorite color is {}", answer);
}

fn current_favorite_color() -> String {
    "blue"
}
