// tests1.rs
//
// 測試對於確保您的程式碼執行您認為它應該執行的操作非常重要。可以使用以下命令運行此文件上的測試：rustlings run tests1
//
// 這個測試有問題——請先讓測試編譯通過，然後讓測試成功，再故意讓測試失敗，以熟悉測試操作。
//
// 執行 `rustlings hint tests1` 或使用 `hint` 子命令獲取提示。

// I AM NOT DONE

#[cfg(test)]
mod tests {
    #[test]
    fn you_can_assert() {
        assert!();
    }
}
