// closures1.rs
// Closures are anonymous functions, but they can be assigned to a variable
// so that they can be used later on
// make me compile and pass

// Execute `rustlings hint closures1` for hints!

// I AM NOT DONE

fn main() {}

#[cfg(test)]
mod tests {
    use super::*;

    # [test]
    fn test_closure() {
        let closure = ;
        assert_eq!("I'm a closure",closure());
    }
}