// generics1.rs
//
// This shopping list program isn't compiling!
// Use your knowledge of generics to fix it.
//
// If you need help, open the corresponding README.md or run: rustlings hint generics1

// I AM NOT DONE

fn main() {
    let mut shopping_list: Vec<?> = Vec::new();
    shopping_list.push("milk");
}
