// functions2.rs
<<<<<<< HEAD
// Make me compile! Execute `rustlings hint functions2` for hints :)s
=======
//
// Execute `rustlings hint functions2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE
>>>>>>> 11d8aea96f2c744d970ed1ffb38785cf5b511e5e

fn main() {
    call_me(3);
}

<<<<<<< HEAD
fn call_me(num: i32) {
=======
fn call_me(num:) {
>>>>>>> 11d8aea96f2c744d970ed1ffb38785cf5b511e5e
    for i in 0..num {
        println!("Ring! Call number {}", i + 1);
    }
}
