// primitive_types2.rs
//
// 填寫剩下缺少代碼的行！沒有提示，沒有技巧，只是讓您習慣打字 :)

// I AM NOT DONE

fn main() {
    // 字符 (`char`)

    // 注意是 _單引號_，這與您經常看到的雙引號不同。
    let my_first_initial = 'C';
    if my_first_initial.is_alphabetic() {
        println!("字母！");
    } else if my_first_initial.is_numeric() {
        println!("數字！");
    } else {
        println!("既不是字母也不是數字！");
    }

    let // 像例子一樣完成這行！您最喜歡的字符是什麼？試試字母、數字、特殊字符，或者不同語言的字符，甚至是表情符號！
    if your_character.is_alphabetic() {
        println!("字母！");
    } else if your_character.is_numeric() {
        println!("數字！");
    } else {
        println!("既不是字母也不是數字！");
    }
}
