// primitive_types3.rs
//
// 創建一個包含至少 100 個元素的陣列，並填寫 ??? 處的代碼。
//
// 執行 `rustlings hint primitive_types3` 或使用 `hint` watch 子命令來獲取提示。

// I AM NOT DONE

fn main() {
    let a = ???

    if a.len() >= 100 {
        println!("哇，那是一個大陣列！");
    } else {
        println!("嗯，我早餐吃的就是這樣的陣列。");
        panic!("陣列不夠大，需要更多元素")
    }
}
